# Assignment — Introduction to Network Science

In this assignment, we will consider some basics of networkx 2.5.

```python
import numpy as np
import networkx as nx
import matplotlib.pyplot as plt
import requests
from zlib import adler32
import pandas as pd
```

```python
assert nx.__version__ == '2.5'
```

In case of an exception, install networkx 2.5 using `pip` as follows:
```
pip uninstall networkx
pip install networkx==2.5
```

### Task 1. Reading a graph from a file (0.5 point)

```python
url = 'https://raw.githubusercontent.com/netspractice/ml-on-graphs/main/datasets/the_lord_of_the_rings_adjlist.txt'
open('the_lord_of_the_rings_adjlist.txt', 'wb').write(requests.get(url).content)

url = 'https://raw.githubusercontent.com/netspractice/ml-on-graphs/main/datasets/the_lord_of_the_rings_edgelist.txt'
open('the_lord_of_the_rings_edgelist.txt', 'wb').write(requests.get(url).content);

```

Usually graph is a sparse data structure. So in most cases data is stored in sparse table formats
1. edge list
2. adjacency list

Edge list is a list of pairs of nodes — a target and destination of edges.

```python
with open('the_lord_of_the_rings_edgelist.txt') as f:
    print(f.read()[:197])
```

Adjacency list is a list of lists of nodes where the first node is a root and others are leafs

```python
with open('the_lord_of_the_rings_adjlist.txt') as f:
    print(f.read()[:190])
```

Most of other specific formats are similar except of syntax and metadata description.

For most cases we can use given functions in networkx:

* `nx.read_adjlist`
* `nx.read_edgelist`
* `nx.read_weighted_edgelist`

In the further task, read two files: `the_lord_of_the_rings_edgelist.txt`, `the_lord_of_the_rings_adjlist.txt` and merge it in a single graph.

<!-- grade_id: cell-3d2796d8ac544282, points: 0, grade: false, solution: true -->
```python
def read_and_merge_graphs(edgelist_path, adjlist_path):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-980a87cb1b6f0dbc, points: 0.5, grade: true, solution: false -->
```python
G = read_and_merge_graphs("the_lord_of_the_rings_edgelist.txt", 
                          "the_lord_of_the_rings_adjlist.txt")

assert nx.number_of_nodes(G) == 35
assert nx.number_of_edges(G) == 98
```

### Task 2. Directed graph (1 point)

The direction in graphs could be crucial feature for many tasks, e.g., link analysis.

We can handle directed graphs with `nx.DiGraph` class. We can use conversion method `nx.to_directed` on the given graph or read it directly passing `nx.DiGraph` constructor to the `nx.read_*` methods.

<!-- grade_id: cell-4ff100c55840771c, points: 0, grade: false, solution: true -->
```python
def convert_to_directed_graph(G):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-678028b3a41d774a, points: 0, grade: false, solution: true -->
```python
def read_directed_graph(path1):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-67792dd08b2f0149, points: 1, grade: true, solution: false -->
```python
assert convert_to_directed_graph(G).is_directed()
G = read_directed_graph("the_lord_of_the_rings_edgelist.txt")
assert G.is_directed()
```

```python
nx.draw(G)
```

### Task 3. Bacis statistics (0.5 point)

Networkx provides a simple interface for getting basic graph statistics:
* `nx.radius`
* `nx.diameter`
* `nx.average_shortest_path_length`
* `nx.average_clustering`

For example, here are statistics of a path graph with 3 nodes:

```python
G = nx.path_graph(3)
print('Radius: ', nx.radius(G))
print('Diameter: ', nx.diameter(G))
print('Average shortest path length: ', nx.average_shortest_path_length(G))
print('Average clustering coefficient: ', nx.average_clustering(G))
nx.draw_networkx(G)
```

Your task is to create graphs by given statistics.

Write a function `graph1` that returns a graph with a radius 3 and diameter 3.

<!-- grade_id: cell-1b71ec27d5c70dc5, points: 0, grade: false, solution: true -->
```python
def graph1():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-02d925c056e1f040, points: 0.2, grade: true, solution: false -->
```python
G = graph1()
assert nx.radius(G) == nx.diameter(G) == 3
```

```python
nx.draw_networkx(G)
```

Write a function `graph2` that returns a graph with a diameter of at least 5 average shortest path lengths.

<!-- grade_id: cell-e6b06880520f2cf7, points: 0, grade: false, solution: true -->
```python
def graph2():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a6c86b83e11199d8, points: 0.2, grade: true, solution: false -->
```python
G = graph2()
assert nx.diameter(G) >= 5 * nx.average_shortest_path_length(G)
```

```python
nx.draw_networkx(G)
```

Write a function `graph3` that returns a graph with an average clustering coefficient between 0.1 and 0.11.

<!-- grade_id: cell-eda0d295af4fc263, points: 0, grade: false, solution: true -->
```python
def graph3():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d4aeb27141847c83, points: 0.1, grade: true, solution: false -->
```python
G = graph3()
assert 0.1 <= nx.average_clustering(G) <= 0.11
```

```python
nx.draw_networkx(G)
```

### Task 4. Graph layout (2 points)

A graph layout is a dictionary with a location of every node: keys are nodes and values are x, y positions. By default in `nx.draw_networkx`, it is the spring layout — a force-directed representation of the network treating edges as springs holding nodes close, while treating nodes as repelling objects.

```python
G = nx.barabasi_albert_graph(10, 2)
layout = nx.spring_layout(G)
layout
```

```python
nx.draw_networkx(G, pos=layout)
```

The Kamada Kawai layout is based on path-length cost-function. Sometimes, it is quite similar to the spring layout, however a significant difference is that the spring layout is random and the Kamada Kawai is deterministic.

```python
layout = nx.kamada_kawai_layout(G)
nx.draw_networkx(G, pos=layout)
```

Circular layout allows to draw a graph where nodes are in a circle.

```python
layout = nx.circular_layout(G)
plt.figure(figsize=(5, 5))
nx.draw_networkx(G, pos=layout)
```

Bipartite layout allows to draw bipartite graph.

```python
B = nx.Graph()
B.add_nodes_from([1, 2, 3, 4], bipartite=0)
B.add_nodes_from(["a", "b", "c"], bipartite=1)
B.add_edges_from([(1, "a"), (1, "b"), (2, "b"), (2, "c"), (3, "c"), (4, "a")])
layout = nx.bipartite_layout(B, [1, 2, 3, 4])
nx.draw_networkx(B, pos=layout)
```

There are other layouts in networkx, also there are some useful parameters that you can find in documentation. Next, let us visualize a coappearance network of characters in The Matrix — a 1999 American science fiction action film written and directed by the Wachowskis. The dataset is taken from https://moviegalaxies.com/.

```python
url = 'https://raw.githubusercontent.com/netspractice/ml-on-graphs/main/datasets/the_matrix_graph.gexf'
open('the_matrix_graph.gexf', 'wb').write(requests.get(url).content);

```

```python
def get_matrix_graph():
    G = nx.read_gexf('the_matrix_graph.gexf')
    mapping = {}
    for node in G.nodes:
        mapping[node] = G.nodes[node]['label']
    G = nx.relabel_nodes(G, mapping)
    return G
matrixG = get_matrix_graph()
plt.figure(figsize=(13, 7))
layout = nx.kamada_kawai_layout(matrixG)
nx.draw_networkx(matrixG, pos=layout)
```

Your task is to draw the matrix graph using the kamada kawai layout with nonlinear modification to obtain a denser layout away from the NEO node. Let us define a proximity coefficient using Gaussian kernel

$$\text{coef}(i) = \exp\left(-\frac{\text{distance}^2}{2\sigma^2}\right)$$

where a _distance_ is the Euclidean distance from the node $i$ to the NEO node, $\sigma$ is the constant. Then the nodes should be moved towards the NEO node such that

$$\text{distance} = \text{coef}(i) \times \text{distance}$$

The simplest way to do it is to center the layout so that the NEO node is in x=0 y=0 and then to multiply all node positions by the proximity coefficient.

Write a function `dense_neo_layout` that takes a layout, sigma and returns a new (dense) layout.

_Hint: use `np.linalg.norm` to calculate Euclidean norm._

<!-- grade_id: cell-6dc95d64261dffa6, points: 0, grade: false, solution: true -->
```python
def dense_neo_layout(layout, sigma):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-fea67b2cc10bdbd6, points: 2, grade: true, solution: false -->
```python
layout = nx.kamada_kawai_layout(matrixG)
layout = dense_neo_layout(layout, sigma=1.5)
assert type(layout) == dict
assert np.all(layout['NEO'] == [0, 0])
assert np.all(layout['REX'].round(2) == [-0.67,  0.58])
assert adler32(str(layout['AGENT BROWN'].round(2)[0]).encode()) == 32309448
```

```python
plt.figure(figsize=(13, 7))
nx.draw_networkx(matrixG, layout)
```

### Task 5. Nodes visualization (0.5 point)

We can set a node size using `draw_networkx` with a parameter `node_size`. It should be a list of values ordered by `G.nodes`. For example, consider a simple path graph.

```python
G = nx.path_graph(3)
G.nodes
```

Let us set a size 200 for the node 0, 300 for 1, 400 for 2.

```python
layout = nx.kamada_kawai_layout(G)
nx.draw_networkx(G, pos=layout, node_size=[200, 300, 400])
```

To color nodes, we can use `node_color`. It should be a list of colors:
* a tuple of float RGB 0-1 values: (red, green, blue)
* a tuple of float RGBA 0-1 values: (red, green, blue, alpha)
* a string — color name

```python
nx.draw_networkx(
    G,
    pos=layout,
    node_color=[(0.4, 0.8, 0, 0.6), (0.7, 0.4, 1, 0.6), (1, 0, 0.5, 0.6)])
```

```python
nx.draw_networkx(
    G,
    pos=layout,
    node_color=['green', 'cyan', 'pink'])
```

Also we can use a colors from matplotlib colormaps. Discrete:

```python
colors = plt.cm.Set1.colors
nx.draw_networkx(
    G,
    pos=layout,
    node_color=[colors[0], colors[1], colors[2]])
```

Or continuous:

```python
rainbow = plt.cm.rainbow
nx.draw_networkx(
    G,
    pos=layout,
    node_color=[rainbow(0.2), rainbow(0.8), rainbow(0.9)])
```

Another option of coloring nodes is to pass a list of single values by the `node_color` and colormap by the `cmap` parameter. In this case, the lowest value will assigned to the first color in colormap and vise versa.

```python
nx.draw_networkx(
    G,
    pos=layout,
    cmap = plt.cm.hot,
    node_color=[1, 2, 3])
```

We can forcibly set a color range using `vmin` and `vmax`.

```python
nx.draw_networkx(
    G,
    pos=layout,
    cmap = plt.cm.hot,
    node_color=[1, 2, 3],
    vmin=0, vmax=4)
```

We can border nodes using `edgecolors` and `linewidths` parameters.

```python
nx.draw_networkx(
    G=G, 
    pos=layout,
    cmap = plt.cm.hot, 
    node_color=[1, 2, 3],
    vmin=0, vmax=4, 
    edgecolors='black', 
    linewidths=1)
```

Your task is to adjust a size and color of nodes in the matrix graph, size and color are taken from a node degree and node attribute `G.nodes[node]`.

```python
matrixG.nodes['NEO']
```

Write a function `drawing_nodes_params` that takes a graph and returns a dictionary with parameters for `nx.draw_networkx` of the form `{'G': graph, 'node_color': color, ...}`. Parameters should contain a graph, kamada kawai layout, node colors from `viz.color` attribute. The colors should be converted from RGB range 0-255 to RGB range 0-1. The node size is a node degree `G.degree[node]` multiplied by 100.

<!-- grade_id: cell-23d41af4fdb48d3c, points: 0, grade: false, solution: true -->
```python
def drawing_nodes_params(graph):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ce378d189964cb4f, points: 0.25, grade: true, solution: false -->
```python
matrixG = get_matrix_graph()
params = drawing_nodes_params(matrixG)
assert type(params) == dict
keys = ['G', 'pos', 'node_color', 'node_size']
assert len(set(params.keys()).intersection(keys)) == 4
assert np.all(params['node_color'][10].round(2) == [0.37, 0.79, 0.75])
assert adler32(str(params['node_color'][0][1].round(2)).encode()) == 19071123
assert params['node_size'][0] == 200
assert params['node_size'][1] == 600
assert nx.to_numpy_array(params['G']).sum() == 326.0
```

```python
plt.figure(figsize=(13, 7))
nx.draw_networkx(**params )
```

Sometimes, it will be useful to draw nodes in nonlinear scale. If we want to make node sizes more comparable, we can use polynomial or exponential scale. On the other hand, we can use logarithmic scale to make sizes more similar.

Write a function `drawing_nodes_polyparams` that returns the same dictionary as the `drawing_nodes_params` function, but the node size is a sum of node degree in a power 3 and 10.

$$\text{size} = \text{degree}^3 + 10$$

<!-- grade_id: cell-7056cc638b9de958, points: 0, grade: false, solution: true -->
```python
def drawing_nodes_polyparams(graph):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6b4ea9133d917597, points: 0.25, grade: true, solution: false -->
```python
matrixG = get_matrix_graph()
params = drawing_nodes_polyparams(matrixG)
assert type(params) == dict
keys = ['G', 'pos', 'node_color', 'node_size']
assert len(set(params.keys()).intersection(keys)) == 4
assert np.all(params['node_color'][10].round(2) == [0.37, 0.79, 0.75])
assert adler32(str(params['node_color'][0][1].round(2)).encode()) == 19071123
assert params['node_size'][0] == 18
assert params['node_size'][1] == 226
assert nx.to_numpy_array(params['G']).sum() == 326.0
```

```python
plt.figure(figsize=(13, 7))
nx.draw_networkx(**params )
```

### Task 6. Edges visualization, weighted graphs (1 points)

We can set an edge width using `draw_networkx` with a parameter `width`. It should be a list of values sorted by G.edges. For example, consider a simple path graph.

```python
G = nx.path_graph(4)
G.edges
```

Let us set a width 1 for the edge 0-1, 5 for 1-2, 10 for 2-3.

```python
layout = nx.kamada_kawai_layout(G)
nx.draw_networkx(G, pos=layout, width=[1, 5, 10])
```

If we create a graph by a weighted ajacency matrix, weights will be stored in an edge attribute `weight`.

```python
W = np.array([[0, 1, 0, 0], 
              [1, 0, 5, 0], 
              [0, 5, 0, 10], 
              [0, 0, 10, 0]])
G = nx.from_numpy_array(W)
```

In this case, we can get a value using `G.edges[edge]` view

```python
G.edges[(0, 1)]
```

Or we can get all values using `nx.get_edge_attributes`

```python
nx.get_edge_attributes(G, 'weight')
```

```python
weights = list(nx.get_edge_attributes(G, 'weight').values())
nx.draw_networkx(G, pos=layout, width=weights)
```

If we want to apply specific color or style to some edges, it is possible by `edge_color` and `style` parameters. For example, we can draw an edge with a weight 10 as a red dotted line.

```python
nx.draw_networkx(
    G, pos=layout, 
    width=[1, 5, 10], 
    edge_color=['black', 'black', 'red'], 
    style=['solid', 'solid', 'dotted'])
```

Your task is to draw the matrix graph where an edge width is 0 if a weight is 1, otherwise it is a weight value itself. An edge color is `red` if weight is at least 5 and `black` otherwise.

<!-- grade_id: cell-2497d776bb6eb00c, points: 0, grade: false, solution: true -->
```python
def drawing_edges_params(graph):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b220499d1d3b12af, points: 1, grade: true, solution: false -->
```python
matrixG = get_matrix_graph()
params = drawing_edges_params(matrixG)
assert type(params) == dict
keys = ['G', 'pos', 'edge_color', 'width']
assert len(set(params.keys()).intersection(keys)) == 4
assert nx.to_numpy_array(params['G']).sum() == 326.0
params['width'][0] == 2.0
assert np.array(params['width']) @ params['width'] == 508
assert adler32(str(params['edge_color']).encode()) == 646700613
```

```python
plt.figure(figsize=(13, 7))
nx.draw_networkx(**params)
```

### Task 7. Power law CDF (0.5 point)

The PDF of the Power law distribution is 

$$ p(x) = Cx^{-\alpha},$$ 

where $C$ is a normalization constant 

$$C = \frac{\alpha - 1}{x_{\text{min}}^{-\alpha + 1}}$$ 

and $\alpha>1$ is called an exponent of the distribution. 

```python
def power_law_pdf(x, alpha=3.5, x_min=1):
    C = (alpha - 1) / x_min ** (1 - alpha)
    return C * x ** (-alpha)
```

```python
x_space = np.linspace(1, 10, 100)
plt.plot(x_space, power_law_pdf(x_space))
plt.title('Power Law PDF');
```

Let us generate observations from a Power Law RV.

The first step is to derive CDF of Powel law: $F(x) = P(X \leq x)$

$$F(x) = 1 - \int_{x}^\infty p(t) dt.$$

Take the integral, derive CDF analytically and write a function `power_law_cdf` with parameters `x`, `alpha` and `x_min`.

<!-- grade_id: cell-0ac914d7d506c37c, points: 0, grade: false, solution: true -->
```python
def power_law_cdf(x, alpha=3.5, x_min=1):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-787e77754a25d952, points: 0.5, grade: true, solution: false -->
```python
assert power_law_cdf(2, 2, 1) == 0.5
assert power_law_cdf(10, 2, 1) == 0.9
assert adler32(str(power_law_cdf(2, 3, 1)).encode()) == 32571595
```

```python
x_space = np.linspace(1, 10, 100)
plt.plot(x_space, power_law_cdf(x_space))
plt.title('Power Law CDF');
```

### Task 8. Power law PPF (0.5 point)

Let $X \sim \text{Power law}$. Next, define a random variable $R$, s.t. $R = F(X)$, so $R$ will be uniformly distributed on interval [0, 1] ([proof](https://en.wikipedia.org/wiki/Probability_integral_transform#Proof)). Good thing here is that we easily can generate uniformly distributed pseudorandom numbers and then transform them into Power Law. Let us find an expression for $x = F^{-1}(r)$, where $r$ is an observation from uniform distrubution on interval [0, 1]. 

Find an analytical form of $F^{-1}(r)$ and write a function `power_law_ppf` (percent point function, also known as a quantile) with parameters `r`, `alpha` and `x_min`. 

<!-- grade_id: cell-3aaaaacf38a1e002, points: 0, grade: false, solution: true -->
```python
def power_law_ppf(r, alpha=3.5, x_min=1):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-03d23ce5a583c2c7, points: 0.5, grade: true, solution: false -->
```python
assert round(power_law_ppf(0.5, 2, 1), 2) == 2
assert round(power_law_ppf(0.9, 2, 1), 2) == 10
assert adler32(str(round(power_law_ppf(0.96, 3, 1), 2)).encode()) == 19792020
```

```python
x_space = np.linspace(0, 0.999, 100)
plt.plot(x_space, power_law_ppf(x_space))
plt.title('Power Law PPF');
```

Finally, we can generate observation from Power law distribution as follows:
1. Generate observation from uniform distribution on interval [0, 1]
2. Calculate PPF value of given observation

```python
def power_law_generate(n, alpha=3.5, x_min=1, random_seed=1):
    np.random.seed(random_seed)
    uni_sample = np.random.uniform(0, 0.999, n)
    return power_law_ppf(uni_sample, alpha, x_min)
```

Look at the histogram of the generated sample

```python
alpha = 3.5
x_min = 1
x_train = power_law_generate(1000, alpha, x_min)
x_space = np.linspace(1, 15, 100)
```

```python
hist, bin_edges = np.histogram(x_train, bins=200, density=True)
bin_centers = (bin_edges[1:] + bin_edges[:-1]) / 2
plt.scatter(bin_centers[hist > 0], hist[hist > 0], s=10)
plt.plot(x_space, power_law_pdf(x_space, alpha, x_min), 
         label='Theoretical PDF', c='tab:orange')
plt.legend()
plt.xlim(1, 10)
plt.show()
```

And here is the same histogram in log-log scale

```python
hist, bin_edges = np.histogram(x_train, bins=200, density=True)
bin_centers = (bin_edges[1:] + bin_edges[:-1]) / 2
plt.scatter(bin_centers, hist, s=10)
plt.plot(x_space, power_law_pdf(x_space, alpha, x_min), 
         label='Theoretical PDF', c='tab:orange')
plt.legend()
plt.xscale('log')
plt.yscale('log')
plt.xlim(1, 10)
plt.show()
```

### Task 9. Estimation of alpha with linear binning (0.5 point)

Given observations from the Power Law distribution, try to estimate $\alpha$. The easiest way is to draw an empirical PDF with linear binning in log-log scale and apply linear regression. By _linear binning_ we mean to keep a bin width is fixed.

Write a function `alpha_lin_bins` that takes a train set, number of linear bins and returns an estimated $\alpha$.

_Hints:_
* _Take log in both side of $p(x) = Cx^{-\alpha}$_
* _To calculate an empirical PDF, use `np.histogram(x_train, bins=bins, density=True)`_
* _To calculate pseudoinverse matrix, use `np.linalg.pinv`_
* _Also you can use fit-predict in `sklearn.linear.LinearRegression`_

<!-- grade_id: cell-64da1436de482e0c, points: 0, grade: false, solution: true -->
```python
def alpha_lin_bins(x_train, bins):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e96f0fe52fe106f4, points: 0.5, grade: true, solution: false -->
```python
bins = 100
x_train = power_law_generate(1000, x_min=1, alpha=2)
assert np.abs(alpha_lin_bins(x_train, bins) - 2) < 0.7
x_train = power_law_generate(1000, x_min=1, alpha=3.5)
assert np.abs(alpha_lin_bins(x_train, bins) - 3.5) < 0.9
x_train = power_law_generate(1000, x_min=1, alpha=10)
assert np.abs(alpha_lin_bins(x_train, bins) - 10) < 2
```

```python
x_min = 1
alpha = 3.5
bins = 100

x_train = power_law_generate(1000, alpha, x_min)
hist, bin_edges = np.histogram(x_train, bins=bins, density=True)
bin_centers = (bin_edges[1:] + bin_edges[:-1]) / 2
plt.scatter(bin_centers, hist, s=10)

hat_alpha = alpha_lin_bins(x_train, bins)
x_space = np.linspace(x_min, x_train.max(), 100)
plt.plot(x_space, power_law_pdf(x_space, hat_alpha, x_min), 
         label='Estimated PDF', c='tab:orange')
plt.legend()
plt.title('Truth alpha = {:.2f}, estimated alpha = {:.2f}'.format(alpha, hat_alpha))
plt.xscale('log')
plt.yscale('log')
plt.show()
```

### Task 10. Degree distribution (1 point)

In this task we will generate different types of graph from its degree distributions.

__Graph A__ should be described with the density `[0, 2, 10]`.

```python
dens_a = [0, 2, 10]
plt.bar(range(len(dens_a)), dens_a);
```

<!-- grade_id: cell-558f01404c6bb378, points: 0, grade: false, solution: true -->
```python
def generate_graph_a() -> nx.Graph:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-75ea8f7cb0526be0, points: 0.2, grade: true, solution: false -->
```python
assert nx.degree_histogram(generate_graph_a()) == dens_a
nx.draw(generate_graph_a())
```

__Graph B__

```python
dens_b = [0, 0, 10]
plt.bar(range(len(dens_b)), dens_b);
```

<!-- grade_id: cell-46de3f5521fe2543, points: 0, grade: false, solution: true -->
```python
def generate_graph_b() -> nx.Graph:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f740d1d577effff7, points: 0.2, grade: true, solution: false -->
```python
assert nx.degree_histogram(generate_graph_b()) == dens_b
nx.draw(generate_graph_b())
```

__Graph C__

```python
dens_c = [0, 0, 0, 0, 5]
plt.bar(range(len(dens_c)), dens_c);
```

<!-- grade_id: cell-81b7d37a1b71dfd6, points: 0, grade: false, solution: true -->
```python
def generate_graph_c() -> nx.Graph:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-248e9156e92d8adb, points: 0.3, grade: true, solution: false -->
```python
assert nx.degree_histogram(generate_graph_c()) == dens_c
nx.draw(generate_graph_c())
```

__Graph D__

```python
dens_d = [0, 5, 0, 0, 0, 1]
plt.bar(range(len(dens_d)), dens_d);
```

<!-- grade_id: cell-072a7c0f71db04ef, points: 0, grade: false, solution: true -->
```python
def generate_graph_d() -> nx.Graph:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f5228d081ca8dd84, points: 0.3, grade: true, solution: false -->
```python
assert nx.degree_histogram(generate_graph_d()) == dens_d
nx.draw(generate_graph_d())
```

### Task 11. MLE of alpha in Power Law (2 points)

Let us estimate $\alpha$ and $x_\min$ of a real social network using maximum likelihood. Here is a real degree distribution of the network of emails between members of the research institution with 1005 nodes.m

```python
url = 'https://raw.githubusercontent.com/netspractice/ml-on-graphs/main/datasets/email_degree_sequence.txt'
open('email_degree_sequence.txt', 'wb').write(requests.get(url).content);

```

```python
degree_sequence = np.loadtxt('email_degree_sequence.txt')
hist, bin_edges = np.histogram(degree_sequence, bins=1000, density=True)
bin_centers = (bin_edges[1:] + bin_edges[:-1]) / 2
plt.scatter(bin_centers[hist > 0], hist[hist > 0], s=5)
plt.title('Email network degree distribution')
plt.xscale('log')
plt.yscale('log')
```

The MLE consists of:
1. Fix $x_\min$ as a minimal node degree (drop node degrees that less than $x_\min$)
2. Calculate $\alpha$ via maximum likelihood estimation using fixed $x_\min$
$$\alpha = 1 + n \left[\sum_i \log \frac{x_i}{x_\min} \right]^{-1}$$
3. Calculate Kolmogorov-Smirnov test
4. Fix $x_\min$ as the next node degree
5. Repeat 2-4 by scanning all possible $x_\min$ and find the best $\alpha$ and $x_\min$ with respect to Kolmogorov-Smirnov test

Write a function `ml_power_law_params` that takes a node degree sequence `degree_sequence` and returns a tuple of two values: the best $\alpha$ and $x_\min$.

_Hints:_
* _to calculate Kolmogorov-Smirnov distance use `scipy.stats.kstest` with a parameter `cdf=power_law_cdf, args=(alpha, x_min)`_
* _since we only consider graphs, check only integer x_min_

<!-- grade_id: cell-9ce569e5194aca18, points: 0, grade: false, solution: true -->
```python
def mle_power_law_params(degree_sequence):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-9bda592a9e1b91da, points: 2, grade: true, solution: false -->
```python
alpha, x_min = 3.5, 3
hat_alpha, hat_x_min = mle_power_law_params(power_law_generate(1000, alpha, x_min))
assert np.abs(hat_alpha - alpha) < 0.05
assert np.abs(hat_x_min - x_min) < 0.001

alpha, x_min = 7, 1
hat_alpha, hat_x_min = mle_power_law_params(power_law_generate(1000, alpha, x_min))
assert np.abs(hat_alpha - alpha) < 0.1
assert np.abs(hat_x_min - x_min) < 0.0001

hat_alpha, hat_x_min = mle_power_law_params(degree_sequence)
assert np.abs(hat_alpha - 2.9) < 0.1
assert np.abs(hat_x_min - 42) < 1
```

```python
hist, bin_edges = np.histogram(degree_sequence, bins=1000, density=True)
bin_centers = (bin_edges[1:] + bin_edges[:-1]) / 2
plt.scatter(bin_centers[hist > 0], hist[hist > 0], s=5)
plt.title('Email network degree distribution')

hat_alpha, hat_x_min = mle_power_law_params(degree_sequence)
x_space = np.linspace(hat_x_min, degree_sequence.max(), 100)
plt.plot(x_space, power_law_pdf(x_space, hat_alpha, hat_x_min), 
         label='Estimated PDF', c='tab:orange')
plt.xscale('log')
plt.yscale('log')
plt.ylim(0.001, 0.5);
```

```python

```